```python
%matplotlib inline
```

```python
from nose.tools import *
# Write your imports in the cell below
```

<!-- grade_id: cell-3a5c071a139ade21, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

from sklearn.model_selection import train_test_split, GridSearchCV

from sklearn.linear_model import LogisticRegression

from sklearn.metrics import f1_score, make_scorer
```

```python
np.random.seed(1234)
```

# Model Training and Improvement Lab
## Comparing and selecting models

### 1. Read the data (1 point)
Like in the previous lab, you need to read the Portuguese bank dataset [here](https://archive.ics.uci.edu/ml/machine-learning-databases/00222/). It has been provided for you in the `data` folder.

Read the dataset using `pandas` (you can use the library with the alias `pd`). Save it in the `bank_data` variable.

<!-- grade_id: cell-39fe3843121d7b36, points: 0, grade: false, solution: true -->
```python
bank_data = pd.read_csv("data/bank.csv", sep=';')
bank_data
```

Output:
```
age            job  marital  education default  balance housing loan  \
0      30     unemployed  married    primary      no     1787      no   no   
1      33       services  married  secondary      no     4789     yes  yes   
2      35     management   single   tertiary      no     1350     yes   no   
3      30     management  married   tertiary      no     1476     yes  yes   
4      59    blue-collar  married  secondary      no        0     yes   no   
...   ...            ...      ...        ...     ...      ...     ...  ...   
4516   33       services  married  secondary      no     -333     yes   no   
4517   57  self-employed  married   tertiary     yes    -3313     yes  yes   
4518   57     technician  married  secondary      no      295      no   no   
4519   28    blue-collar  married  secondary      no     1137      no   no   
4520   44   entrepreneur   single   tertiary      no     1136     yes  yes   

       contact  day month  duration  campaign  pdays  previous poutco
```

```python
bank_data.shape
```

Output:
```
(4521, 17)
```

```python
bank_data.dtypes
```

Output:
```
age           int64
job          object
marital      object
education    object
default      object
balance       int64
housing      object
loan         object
contact      object
day           int64
month        object
duration      int64
campaign      int64
pdays         int64
previous      int64
poutcome     object
y            object
dtype: object
```

<!-- grade_id: cell-b183f83b98f16a29, points: 1, grade: true, solution: false -->
```python
# From now on, all test cells might contain hidden tests. If you follow the instructions correctly, 
# your solution will be graded with maximum points
assert_is_not_none(bank_data)
```

### 2. Preprocess the data (1 point)
Separate explanatory features from labels. Save all features (16 columns total) in the variable `bank_features`. Save the labels (corresponding to the `y` column) in the `bank_labels` variable. Rewrite the labels to be `0` and `1` instead of `no` and `yes`: `bank_labels` should be a numeric column.

<!-- grade_id: cell-fb1c7a38947afc06, points: 0, grade: false, solution: true -->
```python
bank_features, bank_labels = bank_data.drop('y', axis=1), bank_data.y
print(bank_features.shape)
print(bank_labels.shape)
print (f"Count of 'no' values: {len(bank_labels[bank_labels == 'no'])},\nCount of 'yes' values {len(bank_labels[bank_labels == 'yes'])}")
bank_labels = bank_labels.replace(to_replace=['no', 'yes'], value=[0, 1])
print(f"Unique values:{bank_labels.unique()}")
print (f"Count of 0 values: {len(bank_labels[bank_labels == 0])},\nCount of 1 values {len(bank_labels[bank_labels == 1])}")
```

Output:
```
(4521, 16)
(4521,)
Count of 'no' values: 4000,
Count of 'yes' values 521
Unique values:[0 1]
Count of 0 values: 4000,
Count of 1 values 521
```

<!-- grade_id: cell-30a215c478663fb9, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_features)
assert_is_not_none(bank_labels)
```

### 3. Get indicator variables (1 point)
Get indicator (dummy) variables for all categorical columns in `bank_features`. Overwrite the `bank_features` variable to store the new data.

<!-- grade_id: cell-8724b6857c2eb7fc, points: 0, grade: false, solution: true -->
```python
bank_features = pd.get_dummies(bank_features)
bank_features.shape
```

Output:
```
(4521, 51)
```

<!-- grade_id: cell-c2b2c0bd2d30ea5e, points: 1, grade: true, solution: false -->
```python
assert_equal(bank_features.shape, (4521, 51))
```

### 4. Split the data (1 point)
Split the data into training and testing set, with 70% of the data for training. Because the output labels are not equaly distributed, use stratification based on the `bank_labels`.

<!-- grade_id: cell-f46c32bb67654092, points: 0, grade: false, solution: true -->
```python
bank_features_train, bank_features_test, bank_labels_train, bank_labels_test = train_test_split(bank_features, bank_labels, test_size=0.3, stratify=bank_labels, random_state=12)
```

```python
bank_features_train.shape, bank_features_test.shape
```

Output:
```
((3164, 51), (1357, 51))
```

```python
bank_labels_train.shape, bank_labels_test.shape
```

Output:
```
((3164,), (1357,))
```

<!-- grade_id: cell-b14f41fb226c0f37, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_features_train)
assert_is_not_none(bank_labels_train)
assert_is_not_none(bank_features_test)
assert_is_not_none(bank_labels_test)
```

### 5. Train a baseline algorithm (1 point)
Train a logistic regression using the training data. Use 1 000 000 (`1e6`) as the value of C. Score it using the testing data. Save the score in the `baseline_score` variable. You should see a fairly high score.

<!-- grade_id: cell-2fd7633e6266fcba, points: 0, grade: false, solution: true -->
```python
model = LogisticRegression(C=1e6)
model.fit(bank_features_train, bank_labels_train)
baseline_score = model.score(bank_features_test, bank_labels_test)
baseline_score
```

Output:
```
/home/petar/anaconda3/lib/python3.8/site-packages/sklearn/linear_model/_logistic.py:762: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regression
  n_iter_i = _check_optimize_result(
0.89240972733972
```

<!-- grade_id: cell-a3af8c95790cccc4, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(model)
assert_greater(baseline_score, 0.7)
```

### 6. Select a better score (2 points)
As you alrady saw, the positive examples are very few. If you aren't convinced, just check the counts.

We know that the default scoring (accuracy) isn't correct in this case. Better measures would be precision and recall. However, we only want one number. Evaluate the algorithm once again, using a standard scoring method which combines precision and recall. Overwrite the `baseline_score` variable.

Don't forget to score the model on the testing data only.

<!-- grade_id: cell-9b44f24da39ba641, points: 0, grade: false, solution: true -->
```python
def get_baseline_score(y, y_predicted):
    return f1_score(y, y_predicted)
baseline_score = get_baseline_score(bank_labels_test, model.predict(bank_features_test))
print(baseline_score)
```

Output:
```
0.32407407407407407
```

<!-- grade_id: cell-9f4aca08f036d9de, points: 2, grade: true, solution: false -->
```python
assert_less(baseline_score, 0.7)
```

### 7. Tune your model (2 points)
Fine-tune the `C` and `max_iter` parameters.

Use full grid search with the following values:
* `C`: 0.0001, 0.01, 0.1, 1, 10, 100, 10000
* `max_iter`: 50, 100, 300, 1000
* `fit_intercept`: True, False

Save the grid search result in the `grid_search` variable. Don't forget to use the better scoring model that you obtained in the previous task.

<!-- grade_id: cell-88cb1341dfd72476, points: 0, grade: false, solution: true -->
```python
params = {
    'C': [0.0001, 0.01, 0.1, 1, 10, 100, 10000],
    'fit_intercept': [True, False],    
    'max_iter': [50, 100, 300, 1000]
}
f1 = make_scorer(f1_score)
grid_search = GridSearchCV(estimator=LogisticRegression(), param_grid=params, scoring=f1)
grid_search.fit(bank_features_train, bank_labels_train)
```

Output:
```
/home/petar/anaconda3/lib/python3.8/site-packages/sklearn/linear_model/_logistic.py:762: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regression
  n_iter_i = _check_optimize_result(
/home/petar/anaconda3/lib/python3.8/site-packages/sklearn/linear_model/_logistic.py:762: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regres
```

```python
LogisticRegression().get_params().keys()
```

Output:
```
dict_keys(['C', 'class_weight', 'dual', 'fit_intercept', 'intercept_scaling', 'l1_ratio', 'max_iter', 'multi_class', 'n_jobs', 'penalty', 'random_state', 'solver', 'tol', 'verbose', 'warm_start'])
```

```python
grid_search.best_estimator_
```

Output:
```
LogisticRegression(C=10, fit_intercept=False, max_iter=1000)
```

```python
grid_search.cv_results_
```

Output:
```
{'mean_fit_time': array([0.03059082, 0.05842061, 0.1455689 , 0.15162115, 0.02897444,
        0.04247994, 0.03550267, 0.03391843, 0.02942505, 0.05524335,
        0.15859199, 0.4988975 , 0.02910233, 0.05773168, 0.1453886 ,
        0.17769094, 0.02934575, 0.05474072, 0.15660534, 0.49586806,
        0.0301146 , 0.05735779, 0.15389194, 0.4676393 , 0.02962842,
        0.05693846, 0.15739098, 0.49461904, 0.02857318, 0.05789118,
        0.14767628, 0.48035245, 0.02966499, 0.05817618, 0.15459991,
        0.5120626 , 0.02837033, 0.06005611, 0.1499671 , 0.48898654,
        0.02944846, 0.06109834, 0.15445375, 0.50818291, 0.02857261,
        0.05725455, 0.14751873, 0.48446131, 0.02936649, 0.05872936,
        0.15130997, 0.49665322, 0.02853508, 0.05878916, 0.15104232,
        0.73848438]),
 'std_fit_time': array([0.00054315, 0.00861372, 0.0232927 , 0.04616778, 0.00042982,
        0.01159126, 0.00573011, 0.00480755, 0.00066963, 0.0059574 ,
        0.01062352, 0.00820207, 0.00110689, 0.00657899, 0.014
```

<!-- grade_id: cell-0e7311dbe13cf894, points: 2, grade: true, solution: false -->
```python
assert_is_not_none(grid_search)
assert_is_not_none(grid_search.best_estimator_)
```

### 8. Compare scores (1 point)
Use the best estimator from your grid search. Score it using the function from problem 6. Save your answer in `tuned_score`.

<!-- grade_id: cell-98568f584f6c0778, points: 0, grade: false, solution: true -->
```python
tuned_score = get_baseline_score(bank_labels_test, grid_search.best_estimator_.predict(bank_features_test))
```

<!-- grade_id: cell-545b4d41fef4eeef, points: 1, grade: true, solution: false -->
```python
print(tuned_score)
```

Output:
```
0.4297520661157025
```

```python
print(baseline_score - tuned_score)
```

Output:
```
-0.10567799204162842
```

Hmmmm, it seems we have not obtained a better algorithm, even the opposite (the difference is marginal and depends on the random initialization of the cross-validation datasets).

We can, of course, do a lot more things to improve our model's performance, such as normalizing the data, feature selection and feature engineering, trying out different aspects, e.g. polynomial terms, RANSAC; even boosting (we'll talk about this later). However, we'll stop at this point.

What can we conclude? It seems that this is close to the best performance we can get out of this algorithm, given these data points.

We can try improving (cleaning) our dataset, selecting features, etc. but we most likely need a better algorithm. In the next labs, we're going to explore that.